```python
%matplotlib inline
import numpy
import matplotlib.pyplot as plt
```

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel $\rightarrow$ Restart) and then run all cells (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name and collaborators below:

# HW 3:  Interpolation and applications

## Question 1

Consider data at three points $(x_0, y_0) = (0, 0)$, $(x_1, y_1) = (1, 2)$, and $(x_2, y_2) = (2, 2)$.

**(a)** [ 8 pts] Analytically find coefficients of the interpolating  polynomial $P_2(x)$ that passes through these three points in the following bases
1. Monomial: $P(x) = p_0 + p_1 x + p_2 x^2$
1. Lagrange: $P(x) = \sum^2_{k=0} y_k \ell_k(x)$


and show that, independent of basis it's always the same polynomial.



<!-- grade_id: cell-3823725410153810, points: 8, grade: true, solution: true -->
### Problem:

We are given three points:
- $ (x_0, y_0) = (0, 0) $
- $ (x_1, y_1) = (1, 2) $
- $ (x_2, y_2) = (2, 2) $

We need to find the **interpolating polynomial** $ P_2(x) $ that passes through these three points, in two different bases:
1. **Monomial basis**: $ P(x) = p_0 + p_1 x + p_2 x^2 $
2. **Lagrange basis**: $ P(x) = \sum_{k=0}^2 y_k \ell_k(x) $

### 1. Monomial Basis

We start by assuming the polynomial in the **monomial basis** is of the form:

$$
P(x) = p_0 + p_1 x + p_2 x^2
$$

We need to find the coefficients $ p_0 $, $ p_1 $, and $ p_2 $ by using the given points $ (x_0, y_0) $, $ (x_1, y_1) $, and $ (x_2, y_2) $.

#### Step 1: Use the points to form equations

We plug in the points to form three equations.

- For $ (x_0, y_0) = (0, 0) $:
  $$
  P(0) = p_0 + p_1 (0) + p_2 (0)^2 = 0
  $$
  Simplifies to:
  $$
  p_0 = 0
  $$

- For $ (x_1, y_1) = (1, 2) $:
  $$
  P(1) = p_0 + p_1 (1) + p_2 (1)^2 = 2
  $$
  Substituting $ p_0 = 0 $:
  $$
  0 + p_1 + p_2 = 2
  $$
  This simplifies to:
  $$
  p_1 + p_2 = 2 \quad \text{(Equation 1)}
  $$

- For $ (x_2, y_2) = (2, 2) $:
  $$
  P(2) = p_0 + p_1 (2) + p_2 (2)^2 = 2
  $$
  Substituting $ p_0 = 0 $:
  $$
  0 + 2p_1 + 4p_2 = 2
  $$
  This simplifies to:
  $$
  2p_1 + 4p_2 = 2 \quad \text{(Equation 2)}
  $$

#### Step 2: Solve the system of equations

We now have two equations:
1. $ p_1 + p_2 = 2 $
2. $ 2p_1 + 4p_2 = 2 $

- Solve Equation 1 for $ p_1 $:
  $$
  p_1 = 2 - p_2
  $$

- Substitute into Equation 2:
  $$
  2(2 - p_2) + 4p_2 = 2
  $$
  Simplifying:
  $$
  4 - 2p_2 + 4p_2 = 2
  $$
  $$
  4 + 2p_2 = 2
  $$
  Solving for $ p_2 $:
  $$
  2p_2 = 2 - 4 = -2
  $$
  $$
  p_2 = -1
  $$

- Substitute $ p_2 = -1 $ into $ p_1 = 2 - p_2 $:
  $$
  p_1 = 2 - (-1) = 2 + 1 = 3
  $$

#### Step 3: Final Polynomial in Monomial Basis

The coefficients are:
- $ p_0 = 0 $
- $ p_1 = 3 $
- $ p_2 = -1 $

Thus, the polynomial in the **monomial basis** is:

$$
P(x) = 0 + 3x - x^2 = 3x - x^2
$$

---

### 2. Lagrange Basis

Now, we construct the interpolating polynomial in the **Lagrange basis**. The formula for the Lagrange interpolating polynomial is:

$$
P(x) = \sum_{k=0}^{2} y_k \ell_k(x)
$$

Where $ \ell_k(x) $ is the **Lagrange basis polynomial** for point $ k $, defined as:

$$
\ell_k(x) = \prod_{\substack{i=0 \\ i \neq k}}^{2} \frac{x - x_i}{x_k - x_i}
$$

#### Step 1: Construct $ \ell_0(x) $

We want $ \ell_0(x) $ to be 1 at $ x_0 = 0 $, and 0 at $ x_1 = 1 $ and $ x_2 = 2 $.

Using the formula:

$$
\ell_0(x) = \frac{x - x_1}{x_0 - x_1} \cdot \frac{x - x_2}{x_0 - x_2}
$$

Substitute the values:

$$
\ell_0(x) = \frac{x - 1}{0 - 1} \cdot \frac{x - 2}{0 - 2} = (1 - x) \cdot \frac{2 - x}{2}
$$

Thus:

$$
\ell_0(x) = \frac{(1 - x)(2 - x)}{2}
$$

#### Step 2: Construct $ \ell_1(x) $

We want $ \ell_1(x) $ to be 1 at $ x_1 = 1 $, and 0 at $ x_0 = 0 $ and $ x_2 = 2 $.

Using the formula:

$$
\ell_1(x) = \frac{x - x_0}{x_1 - x_0} \cdot \frac{x - x_2}{x_1 - x_2}
$$

Substitute the values:

$$
\ell_1(x) = \frac{x - 0}{1 - 0} \cdot \frac{x - 2}{1 - 2} = x \cdot (2 - x)
$$

Thus:

$$
\ell_1(x) = x(2 - x)
$$

#### Step 3: Construct $ \ell_2(x) $

We want $ \ell_2(x) $ to be 1 at $ x_2 = 2 $, and 0 at $ x_0 = 0 $ and $ x_1 = 1 $.

Using the formula:

$$
\ell_2(x) = \frac{x - x_0}{x_2 - x_0} \cdot \frac{x - x_1}{x_2 - x_1}
$$

Substitute the values:

$$
\ell_2(x) = \frac{x - 0}{2 - 0} \cdot \frac{x - 1}{2 - 1} = \frac{x}{2} \cdot (x - 1)
$$

Thus:

$$
\ell_2(x) = \frac{x(x - 1)}{2}
$$

#### Step 4: Construct the Final Polynomial

Now that we have the Lagrange basis polynomials, we can construct the full polynomial:

$$
P(x) = y_0 \ell_0(x) + y_1 \ell_1(x) + y_2 \ell_2(x)
$$

Substitute the values $ y_0 = 0 $, $ y_1 = 2 $, and $ y_2 = 2 $:

$$
P(x) = 0 \cdot \ell_0(x) + 2 \cdot \ell_1(x) + 2 \cdot \ell_2(x)
$$

Simplify:

$$
P(x) = 2 \cdot \left[ x(2 - x) \right] + 2 \cdot \left[ \frac{x(x - 1)}{2} \right]
$$

Simplify each term:

$$
P(x) = 2 \cdot (2x - x^2) + x(x - 1)
$$

$$
P(x) = (4x - 2x^2) + (x^2 - x)
$$

Combine like terms:

$$
P(x) = 4x - x^2 - x
$$

Finally:

$$
P(x) = 3x - x^2
$$

---

### Conclusion:

Whether using the **Monomial basis** or the **Lagrange basis**, the interpolating polynomial is the same:

$$
P(x) = 3x - x^2
$$


**(b)** [4 pts] Any $N+1$ distinct values of $x$ uniquely define the $N+1$ Lagrange polynomials $\ell_0(x), \ell_1(x)\ldots\ell_N(x)$

show that these functions satisfy 

$$\sum^N_{i=0} \ell_i(x) = 1$$

for all values of $x$ (not just the nodes).  We say that the Lagrange polynomials form a "Partition of Unity" because they always sum to 1.

Hint:  consider interpolation of a constant function, $f(x) = c$ sampled at $N+1$ distinct points.


<!-- grade_id: cell-5902170639545808, points: 4, grade: true, solution: true -->
### Problem:

We are asked to show that the sum of the Lagrange basis polynomials $( \ell_0(x), \ell_1(x), \ldots, \ell_N(x) $) satisfies:

$$
\sum_{i=0}^N \ell_i(x) = 1
$$

for **all values of \( x \)**, not just at the interpolation nodes. This is known as the **Partition of Unity** property of Lagrange polynomials.

### Strategy:

To prove this, we will consider interpolating a **constant function** $( f(x) = c $), where $( c $) is a constant, sampled at $( N+1 $) distinct points $( x_0, x_1, \ldots, x_N $). 

Since the function is constant, the Lagrange interpolation should result in the constant function itself. This will help us derive the result that the sum of the Lagrange basis polynomials is 1.

### Lagrange Interpolation of a Constant Function:

The **Lagrange interpolation polynomial** for a function $( f(x) $) at $( N+1 $) distinct points is given by:

$$
P(x) = \sum_{i=0}^N f(x_i) \ell_i(x)
$$

where:
- $( f(x_i) $) is the value of the function at the interpolation point $( x_i $),
- $( \ell_i(x) $) is the Lagrange basis polynomial that is 1 at $( x = x_i $) and 0 at all other interpolation points.

Now, consider the **constant function** $( f(x) = c $). For this function, we have:

$$
f(x_i) = c \quad \text{for all } i
$$

Thus, the interpolation polynomial becomes:

$$
P(x) = \sum_{i=0}^N c \cdot \ell_i(x)
$$

Factor out the constant \( c \):

$$
P(x) = c \cdot \sum_{i=0}^N \ell_i(x)
$$

### Key Insight:

Since $( f(x) = c $) is a constant function, we know that the interpolation polynomial $( P(x) $) must also be equal to $( c $) for all values of $( x $), not just at the interpolation points:

$$
P(x) = c \quad \text{for all } x
$$

Therefore, we can equate this to the previous expression for \( P(x) \):

$$
c \cdot \sum_{i=0}^N \ell_i(x) = c
$$

### Divide by $( c $) (assuming $( c \neq 0 $)):

We now divide both sides by $( c $):

$$
\sum_{i=0}^N \ell_i(x) = 1
$$

This shows that the sum of the Lagrange basis polynomials is 1 for all values of \( x \), not just at the interpolation nodes.

### Conclusion:

We have shown that:

$$
\sum_{i=0}^N \ell_i(x) = 1
$$

for all $( x $). This property is known as the **Partition of Unity**, as the Lagrange polynomials collectively sum to 1 across the entire domain of the interpolation, regardless of the specific points $( x_0, x_1, \ldots, x_N $), as long as they are distinct.


## Question 2: Fun with Chebyshev Polynomials

**(a)** [4 pts] The $n$th Chebyshev polynomial is characterized (up to a constant) by the identity

$$T_n(\cos \theta) = \cos (n \theta)$$

Use this identity to show that the Chebyshev polynomials are orthogonal on $x\in[-1,1]$ with respect to the weight

$$\omega(x) = \frac{1}{\sqrt{1 - x^2}}$$

To do this you must prove that

$$\int^1_{-1} \omega(x) T_n(x) T_m(x) dx = \left \{ \begin{aligned} 
a & ~ & m = n \\ 
0 & ~ & m \neq n 
\end{aligned} \right .$$

where $a$ is a finite constant (also find this coefficient).  Note that you may have multiple non-zero coefficients $a$.  Make sure to cover all these cases and list the conditions that you will find these values.

<!-- grade_id: cell-8354694405440956, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

**(b)** [4 pts] Consider the  Gaussian 
$$ f(x) = e^{-(x-3)^2} $$ on the interval $x\in[0,6]$

make a plot comparing the polynomial interpolant through $N=11$ points using 
* Evenly spaced points
* The appropriate Chebyshev nodes (i.e. roots of $T_{11}$: **note** you will need to transform the nodes defined on the interval $s\in[-1,1]$ to $x\in[0,6]$)

Be sure to show both the interpolation points and a reasonably resolved interpolating polynomial $P_{10}(x)$, and label all lines and axes.

You can use any of the routines from the notes or `numpy.polyfit` and `numpy.polyval`.  
For fun, experiment with other functions (e.g. Runge's function $f(x) = 1/(1 + 25x^2)$)

<!-- grade_id: cell-02c4fc987bf9ddca, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

**(c)** [4 pts] define the relative error as
        
$$
 r = \frac{||f(x) - P_N(x)||}{||f(x)||}
$$

where $||\cdot||$ is an appropriate norm (length) of a vector, in python you can calculate this by

```python
x = numpy.linspace(x_min, x_max)
P_N = numpy.polyval(p, x)

err = f(x) - P_N(x)
r = numpy.linalg.norm(err)/numpy.linalg.norm(f(x))
```

* Make a well labeled plot showing log(r) vs number of grid points $N$ for $N\in[5,30]$ 
* Comment briefly on your results. Some questions to consider are:
    *  does $P_N(x)\rightarrow f(x)$ as $N\rightarrow\infty$?
    *  Is one interpolant more accurate than the other?
    * Anything else you notice

<!-- grade_id: cell-3a61e29e4740bd7d, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-9bd361f6e06eb56c, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3

We can often reformulate finite difference approximations as matrix-vector products.   

Here we will construct a matrix   $D$  such that multiplying a vector of $\mathbf{y}$ values would lead to a second order approximation of the First derivative for the given data. 

In other words, defining the two vectors   $\mathbf{x}=[x_0, x_1,\ldots,x_N]^T$ and $\mathbf{y}=[y_0, y_1,\ldots,y_N]^T$  such that

$$\mathbf{y} = f(\mathbf{x})$$

for some scalar function $f$,  find a matrix $D$ such that 

$$
f'(\mathbf{x}) \approx D\mathbf{y}
$$

**(a)** [16] Using any method of your choice, derive the following finite-difference formulas for 3 equispaced points $x_i, x_{i+1}, x_{i+2}$ 

* 2nd-order accurate approximation to the *1st derivative* at  point 
    * $x_i$ (forward differences)
    * $x_{i+1}$ (centered differences)
    * $x_{i+2}$ (backward differences)
* 2nd-order accurate approximation to the *2nd derivative* evaluated at each point (hint: think before you grind)

Some hints:
1. consider the interpolating polynomial $P_2(x)$ through those 3 points and calculate $P'_2(x)$ at each point.
2. the method of undetermined coefficients on Taylor's series can be useful here as well



<!-- grade_id: cell-88d2d6f7dba4237d, points: 16, grade: true, solution: true -->
YOUR ANSWER HERE

**(b)** [4] Now consider an equispaced set of points  $y_i = f(x_i)$.  Can you describe the structure of a consistent 2nd order, first derivative matrix $D$,  such that 

$$
    f'(\mathbf{x}) \approx D\mathbf{y}
$$

returns the consistent 1st derivative at every point (including the end points)

<!-- grade_id: cell-2125755949784286, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

**(c)** [4] Using the function `fdcoeffV.py` described in class (and included with this homework), write a python function that takes in a numpy array of coordinates $x$ (not necessarily evenly spaced)  and returns a matrix $D$ that calculates the consistent 2nd-order 1st derivatives at each point using 3 nearest neighbors (be careful at the edges)

Given this matrix $D$, it is easy to test it for various functions and meshes

<!-- grade_id: cell-5516025382221234, points: 0, grade: false, solution: true -->
```python
from  fdcoeffV import fdcoeffV

def D1(x):
    """
    Write a good doc-string
    
    x: ndarray of coordinates
    D: ndarray
        Matrix (2-D ndarray ) to calculate consistent 2nd order first derivative
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    

```

<!-- grade_id: cell-7511846809958937, points: 4, grade: true, solution: false -->
```python
# Let's test this using f(x) = sin(pi*x) on a uniform mesh
N = 200
x = numpy.linspace(-1, 1, N)
dx = x[1]-x[0]

pi = numpy.pi
y = lambda x: numpy.sin(pi*x)
y_prime = lambda x: pi*numpy.cos(pi*x)

print("Uniform mesh: ",end='')
numpy.testing.assert_allclose(numpy.dot(D1(x), y(x)), y_prime(x), rtol=1e-2)
print("Success!")

# test on a randomized mesh
x_rand = x + dx/3.*numpy.random.rand(N)
print("Random mesh: ",end ='')
numpy.testing.assert_allclose(numpy.dot(D1(x_rand), y(x_rand)), y_prime(x_rand), rtol=1e-2)
print("Success!")


```

**(d)** [6] make a beautiful convergence plot showing how the relative error in your finite element approximation

$$ r = \frac{||f'(\mathbf{x}) - D\mathbf{y}||}{||f'(\mathbf{x})||}$$

 scales with mesh spacing $\Delta x$ for $f(x) = \sin(\pi x)$ on the interval $x\in[-1,1]$.  What is the expected rate of convergence?  Compare with the best fit straight line (in log space) to your error estimates.

<!-- grade_id: cell-2dac3533cd822755, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cdbaae508b97aeab, points: 2, grade: true, solution: true -->
YOUR ANSWER HERE

**(e)** [4] Suppose we knew the *first derivative* of a function $f(x)$ at a discrete set of points and wanted to solve for the discrete function itself.  i.e.  ideally given $D$ and vector $\mathbf{y} = f'(\mathbf{x})$,  we would like to be able to solve the linear algebra problem

$$
    D\mathbf{f} = \mathbf{y}
$$

for $\mathbf{f}\approx f(\bf{x})$.  Discuss any issues we might have with this approach.

<!-- grade_id: cell-3816704844311840, points: 4, grade: true, solution: true -->
YOUR ANSWER HERE

```python

```